<nav>
    <a href="for_loop_exercises.ipynb">&#9664; Übungsaufgaben zu Schleifen</a> |
    <a href="group_3_level_2_75.ipynb">Viel Unterstützung &#9685;</a> |
    <a href="group_3_level_2.ipynb">Keine Unterstützung &#9711;</a> |
</nav>

<div class="alert alert-warning">
<strong>Aufgabe 1</strong>

Im Rahmen der Materialverwaltung soll ein Programm zur automatischen Berechnung eines Materialbestandes, anhand von Eingängen und Abgängen eines Materials in einer Lagerhalle, implementiert werden.  

Implementieren Sie nun die Funktion `calculate_stock`. Diese Funktion erhält als Parameter eine Liste `incoming_outgoing`, welche in Form von Integer-Werten die Materialeingänge und Abgänge in die Lagerhalle angibt (in etwa `[+500, -200, +70, -65]`).  
Als Ergebnis soll der Materialbestand in Form eines Integer-Wertes zurückgegeben werden ("stock").

</div>

<div class="alert alert-block alert-info"><strong>Hilfe</strong>

Sie können Ihre Lösung mit den automatisierten [Tests](#tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `for`-Schleife in Python finden Sie [hier](../for_loop.ipynb).
</div>

<!-- grade_id: cell-080ae4f875168748, points: 0, grade: false, solution: true -->
```python
def calculate_stock(incoming_outgoing):
    
    ###BEGIN SOLUTION
    stock = 0
    for entry in incoming_outgoing:
        stock += entry
    
    ###END SOLUTION
    return stock

print("Bei \"incoming_outgoing = [20, 30, 40]\" liefert Ihre Funktion:", calculate_stock([20, 30, 40]))
```

Output:
```
Bei "incoming_outgoing = [20, 30, 40]" liefert Ihre Funktion: 90
```

---
<a id="tests"></a>
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-9563226530e72506, points: 0, grade: true, solution: false -->
```python
try:
    calculate_stock([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Berechnung
Der Materialbestand (`stock`) soll für jede Liste korrekt berechnet werden.

<!-- grade_id: cell-e412b7ff90a78f87, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

incoming_outgoing = [+519, -250, -19, -12, +234]
stock = calculate_stock(incoming_outgoing)
try:
    assert_equal(stock, 472)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [+99999999, -1234567, -9876543]
stock = calculate_stock(incoming_outgoing)
try:
    assert_equal(stock, 88888889)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

incoming_outgoing = [-12, +20]
stock = calculate_stock(incoming_outgoing)
try:
    assert_equal(stock, 8)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
```

Output:
```
Test 1 erfolgreich
Test 2 erfolgreich
Test 3 erfolgreich
```