<div class="alert alert-warning">
<strong>Aufgabe 3</strong>

Im Rahmen der Steuerrechnung soll ein Programm zur automatischen Analyse der Umsatzsteuer, welche von Ihrem Unternehmen eingenommen wird, implementiert werden. Dazu sollen die Netto-Verkaufspreise analysiert werden. So soll aus einer Liste mit vielen Einnahmen heruasgefiltert werden, welches die höchste eingenommene Umsatzsteuer ist. Außerdem sollen der durchschnitt sowie die Summe aller eingenommenen Umsatzsteuern berechnet werden.

Implementieren Sie nun die Funktion `analyse_taxes`. Diese Funktion erhält als Parameter eine Liste `sales`, welche in Form von Float-Werten die Netto-Verkaufspreise aller verkauften Artikel angibt (in etwa `[500.00, 299.99, 1.79, 149.49]`).   
Als Ergebnis soll der gesamte Betrag, welcher durch die Umsatzsteuer eingenommen wurde, in Form eines Float-Wertes zurückgegeben werden ("taxes").   
Die durchschnittliche Höhe der Umsatzsteuer ("average").  
Der höchste Betrag, welcher durch die Umsatzsteuer an einem Artikel eingenommen wurde, in Form eines Float-Wertes ("highest_tax").   

Tipp: $ Umsatzsteuer = Nettobetrag * 0.19 $

</div>

<div class="alert alert-block alert-info"><strong>Hilfe</strong>

Sie können Ihre Lösung mit den automatisierten [Tests](#tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `for`-Schleife in Python finden Sie [hier](../for_loop.ipynb).
</div>

<!-- grade_id: cell-ea4effec3a83d25e, points: 0, grade: false, solution: true -->
```python
def analyse_taxes(sales):
    
    ###BEGIN SOLUTION
    taxes = 0
    highest_tax = 0
    for sale in sales:
        taxes += sale *0.19
        if (sale * 0.19) > highest_tax:
            highest_tax = sale * 0.19
            
    average = taxes / len(sales)
    
    ###END SOLUTION
    return taxes, average, highest_tax

print("Bei \"sales = [20, 30, 40]\" liefert Ihre Funktion:", analyse_taxes([20, 30, 40]))
```

Output:
```
Bei "sales = [20, 30, 40]" liefert Ihre Funktion: (17.1, 5.7, 7.6)
```

---
<a id="tests"></a>
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.
- Test 1 - Korrekte Berechnung der gesamten Umsatzsteuer
> Die Umsatzsteuer (`taxes`) soll für jede Liste korrekt berechnet werden.
- Test 2 - Korrekte Berechnung des Durchschnitts
> Der Durchschnitt (`average`) soll für jede Liste korrekt berechnet werden.
- Test 3 - Korrekte Bestimmung der höchsten Umsatzsteuer
> Der Betrag (`highest_tax`) soll für jede Liste korrekt herausgefiltert werden.

<!-- grade_id: cell-8baa31006b25d951, points: 0, grade: true, solution: false -->
```python
import unittest
class TestNotebookClass(unittest.TestCase):
    """Testing function create_production_plan"""
    sales1 = [300, 250.45, 19.78, 123.56, 50.23]
    sales2 = [99999999, 1234567, 9876543]
    sales3 = [12, 20]
    sales4 = [500, 250.45, 19.78, 123.56, 50.23]
    sales5 = [20]
    def test11(self):
        """Test1: Ist taxes korrekt?"""
        out = analyse_taxes(self.sales1)
        expected = 141.3638
        self.assertTrue(out[0], expected)
    def test12(self):
        """Test1: Ist taxes korrekt?"""
        out = analyse_taxes(self.sales2)
        expected = 21111110.71
        self.assertTrue(out[0], expected)
    def test13(self):
        """Test1: Ist taxes korrekt?"""
        out = analyse_taxes(self.sales3)
        expected = 6.08
        self.assertTrue(out[0], expected)
    def test21(self):
        """Test2: Ist average korrekt?"""
        out = analyse_taxes(self.sales4)
        expected = 35.87276
        self.assertTrue(out[2], expected)
    def test22(self):
        """Test2: Ist average korrekt?"""
        out = analyse_taxes(self.sales2)
        expected = 7037036.903333333
        self.assertTrue(out[2], expected)
    def test23(self):
        """Test2: Ist average korrekt?"""
        out = analyse_taxes(self.sales3)
        expected = 3.04
        self.assertTrue(out[2], expected)
    def test24(self):
        """Test2: Ist average korrekt?"""
        out = analyse_taxes(self.sales5)
        expected = 3.8
        self.assertTrue(out[2], expected)
    def test31(self):
        """Test3: Ist highest_tax korrekt?"""
        out = analyse_taxes(self.sales1)
        expected = 57
        self.assertTrue(out[3], expected)
    def test32(self):
        """Test3: Ist highest_tax korrekt?"""
        out = analyse_taxes(self.sales2)
        expected = 18999999.81
        self.assertTrue(out[3], expected)
    def test33(self):
        """Test3: Ist highest_tax korrekt?"""
        out = analyse_taxes(self.sales3)
        expected = 3.8
        self.assertTrue(out[3], expected)
    def test34(self):
        """Test3: Ist highest_tax korrekt?"""
        out = analyse_taxes(self.sales5)
        expected = 3.8
        self.assertTrue(out[3], expected)

import sys
suite = unittest.TestLoader().loadTestsFromTestCase(TestNotebookClass)
unittest.TextTestRunner(verbosity=4,stream=sys.stderr).run(suite)
```