<nav>
    <a href="for_loop_exercises.ipynb">&#9664; Übungsaufgaben zu Schleifen</a> |
    <a href="group_2_level_2_75.ipynb">Viel Unterstützung &#9685;</a> |
    <a href="group_2_level_2.ipynb">Keine Unterstützung &#9711;</a> |
</nav>

<div class="alert alert-warning">
<strong>Aufgabe 2</strong>

Es soll ein Programm zur automatischen Analyse von Kontobewegungen implementiert werden.
So soll aus einer großen Liste von Zugängen und Abgängen der maximale Betragswert erkannt werden.

Implementieren Sie nun die Funktion `analyse_transfers`.
Diese Funktion erhält als Parameter eine Liste `transfers`, welche in Form von Float-Werten die Zugänge und Abgänge auf das Konto angibt (in etwa `[+500.00, -200.00, +0.70, -0.65]`).
Als Ergebnis soll der `höchste Transfer` als Float-Wert zurückgegeben werden ("highest_transfer").
Hierbei sollen sowohl negative als auch positive Transfers berücksichtigt werden.
</div>


<div class="alert alert-block alert-info"><strong>Hilfe</strong>
    
Sie können Ihre Lösung mit den automatisierten [Tests](#tests) weiter unten im Dokument überprüfen.<br>
Eine Einführung in die `for`-Schleife in Python finden Sie [hier](../for_loop.ipynb).
</div>

<!-- grade_id: cell-29bb7125c4142d63, points: 0, grade: false, solution: true -->
```python
def analyse_transfers(transfers):
    
    ###BEGIN SOLUTION
    # initialize max value
    highest_transfer = 0
    for transfer in transfers:
        # compare squared values (alternatively use abs function)
        #if abs(transfer) > abs(highest_transfer):
        if transfer ** 2 > highest_transfer ** 2:
            highest_transfer = transfer
    
    ###END SOLUTION
    return highest_transfer

print("Bei \"transfers = [20, 30, 40]\" liefert Ihre Funktion:", analyse_transfers([20, 30, 40]))
```

---
<a id="tests"></a>
# Automatische Tests
Ab hier finden Sie einige automatisierte Tests um Ihre Lösungen zu überprüfen. Um Ihre Lösung zu prüfen führen Sie bitte zuerst die Zelle mit Ihrer Lösung und danach die Zelle mit dem zugehörigen Test aus.


<!-- grade_id: cell-f1d46b35aa7bdd29, points: 0, grade: true, solution: false -->
```python
try:
    analyse_transfers([0, 0, 0, 0])
except NotImplementedError:
    print("Bitte entfernen Sie das \"raise NotImplementedError()\"")
```

## Test 1 - Korrekte Bestimmung der höchsten Bankbewegung
Die höchste Bankbewegung (`highest_transfer`) soll für jede Liste korrekt bestimmt werden.  
Hierbei gilt es insbesondere auch negative Bewegungen zu berücksichtigen.

<!-- grade_id: cell-0bfb883eacffd309, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

transfers = [+100.45, -26.89, -23.00, +210.00, +234.00]
highest_transfer = analyse_transfers(transfers)
try:
    assert_equal(highest_transfer, 234)
    print("Test 1 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+200.00, -200.00, +200.00, -200.00]
highest_transfer = analyse_transfers(transfers)
try:
    assert_equal(highest_transfer, 200)
    print("Test 2 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+15.00]
highest_transfer = analyse_transfers(transfers)
try:
    assert_equal(highest_transfer, 15)
    print("Test 3 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch")
    print(e)
    

transfers = [+300.00, -500.00, -23.00, +210.00, +234.00]
highest_transfer = analyse_transfers(transfers)
try:
    assert_equal(highest_transfer, -500)
    print("Test 4 erfolgreich")
except AssertionError as e:
    print("Das Ergebnis ist falsch, beachten Sie auch negative Buchungen")
    print(e)
```